# Basic Numpy

In this assignment we will explore some basic numpy assignments. 

We are building an airline routing system, we want to know how many ways there are to reach airport $b$ from airport $a$ in $K$ flights.

Given an $N \times N$ adjacency matrix $M$, write an algorithm that works out how many paths of length $K$ exist from any given node to another. The answer should be represented as a $N \times N$ matrix with $ N_{ab} $ being the amount of paths from node $a$ to node $b$.

This cell can not be edited.

<!-- grade_id: cell-02baf050bea44f2a, points: 0, grade: false, solution: true -->
```python
import numpy as np

def reachable(m, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9c47b3c11ef9d1b0, points: 3, grade: true, solution: false -->
```python
flight_network_m = np.array([[0, 0, 0], [1, 0, 1], [1, 1, 1]])

np.testing.assert_array_equal(
    np.eye(3), reachable(flight_network_m, 0),
    "With 0 steps only the current node should be reachable"
)

np.testing.assert_array_equal(
    flight_network_m, reachable(flight_network_m, 1),
    "One step should equal the input adjacency matrix"
)

np.testing.assert_array_equal(
    np.array([[0,0,0],[2,1,2],[3,2,3]]), reachable(flight_network_m, 3),
    "One step should equal the input adjacency matrix"
)
```

Visualize the graph defined by `flight_network_m` using your favorite plotting library in the cell below

<!-- grade_id: cell-4ba6195442b48c46, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Finally this is a manually graded task worth 1 point, it should not be possible to edit this cell.